<a href="https://colab.research.google.com/github/Pranj9945/-CE880_Lab_coursework-/blob/main/week1/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-1 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).


---
## Question 1:

Write a Python function to calculate simple interest. The formula for calculating simple interest is:

`Simple Interest (SI) = (Principal * Rate * Time) / 100`

#### Parameters:
* principal (float): The principal amount (initial amount of money).
* rate (float): The rate of interest per time period.
* time (float): The time for which the money is borrowed or invested.



<!-- grade_id: myPythagorean_ans, points: 0, grade: false, solution: true -->
```python
import math
def calculate_simple_interest(principal, rate, time):
    """Create a function that takes three parameters (principal, rate, time)
    and returns the calculated simple interest."""
    # YOUR CODE HERE
    simple_interest = (principal * rate * time) / 100
    return simple_interest
    raise NotImplementedError()
```

```python
calculate_simple_interest(1000, 5, 3)
```

Output:
```
150.0
```

<!-- grade_id: myPythagorean_test, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert calculate_simple_interest(1000, 2.22, 3) == 66.6
assert calculate_simple_interest(5000, 0, 3) == 0
assert calculate_simple_interest(5000, 5, 3) == 750
```

---

## Question 2:

Write a Python function to calculate the volume of a cylinder. The formula for calculating the volume of a cylinder is:

`Volume = π * radius^2 * height`

Hint: For simplicity, you can use the value of π as 3.14159.


#### Parameters:
* radius (float): The radius of the cylinder's base.
* height (float): The height of the cylinder.

<!-- grade_id: areaofTrap_ans, points: 0, grade: false, solution: true -->
```python
def calculate_cylinder_volume(radius, height):
    """Write a Python function to calculate the volume of a cylinder."""
    # YOUR CODE HERE
    pi = 3.14159
    volume = pi * (radius ** 2) * height
    return volume
    raise NotImplementedError()
```

```python
calculate_cylinder_volume(10, 4)
```

Output:
```
1256.636
```

<!-- grade_id: areaofTrap_test, points: 1, grade: true, solution: false -->
```python
"""Checkthe correct answer for various inputs."""
assert calculate_cylinder_volume(5, 10) == 785.3975
assert calculate_cylinder_volume(15, 8) == 5654.862
assert calculate_cylinder_volume(20, 3) == 3769.908
```

---
## Question 3:

 Write a Python program to reverse a string. Example: Sample String: `1234xyz` Expected Output: `zyx4321`

<!-- grade_id: reversestring_ans, points: 0, grade: false, solution: true -->
```python
def reverse_string(input_str):
    """Write your code here to reverse the string"""

    # YOUR CODE HERE
    reversed_str = ''
    for i in range(len(input_str)-1, -1, -1):
      reversed_str += input_str[i]
    return reversed_str
    raise NotImplementedError()
```

```python
reverse_string('ChatGPT')
```

Output:
```
'TPGtahC'
```

<!-- grade_id: reversestring_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert reverse_string('Humanitarian') == 'nairatinamuH'
assert reverse_string('ChatGPT') == 'TPGtahC'
assert reverse_string('Data Science') == 'ecneicS ataD'
assert reverse_string('CE880') == '088EC'
```

## Question 4:
Write a Python program to print the odd numbers from a given list and in ascending order.

Sample List : `[1, 2, 3, 4, 5, 6, 7, 8, 9]` Expected Result : `[1, 3, 5, 7, 9]`

<!-- grade_id: even_num_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def odd_num(lst):
    # Filter the odd numbers
    odd_numbers = [num for num in lst if num % 2 != 0]

    # Sort the odd numbers in ascending order
    odd_numbers.sort()

    # Print the sorted odd numbers
    return odd_numbers

    # Call the function with the sample list
    print_odd_numbers(sample_list)

    # Call the function with the sample list
    print_odd_num(sample_list)
    raise NotImplementedError()
```

```python
odd_num([2, 3, 14, 17, 8, 19, 11, 14])
```

Output:
```
[3, 11, 17, 19]
```

<!-- grade_id: odd_num_test, points: 1, grade: true, solution: false -->
```python
"""Check that odd_num returns the correct answer for various inputs."""
assert odd_num([1, 2, 3, 4, 6, 7, 8, 9]) == [1, 3, 7, 9]
assert odd_num([2, 3, 4, 7, 8, 9, 11, 14]) == [3, 7, 9, 11]
assert odd_num([2, 3, 14, 17, 8, 19, 11, 14]) == [3, 11, 17, 19]
```

## Question 5:
Write a Python function to create and print a list where the values are cube of numbers between `1` and `n`, where `n` is user defined number: Example `cube_values(4) == [1, 8, 27, 64]`

<!-- grade_id: sqr_values_ans, points: 0, grade: false, solution: true -->
```python
def cube_values(n):
    """Write a Python program to print the cube of numbers between 1 and `n`"""

    # YOUR CODE HERE
def cube_values(n):
    """Create and print a list where the values are cubes of numbers between 1 and n."""
    cube_list = [i ** 3 for i in range(1, n + 1)]
    return cube_list
```

```python
cube_values(9)
```

Output:
```
[1, 8, 27, 64, 125, 216, 343, 512, 729]
```

<!-- grade_id: sqr_values_test, points: 1, grade: true, solution: false -->
```python
"""Check that cube_values returns the correct answer for various inputs."""
assert cube_values(4) == [1, 8, 27, 64]
assert cube_values(8) == [1, 8, 27, 64, 125, 216, 343, 512]
assert cube_values(9) == [1, 8, 27, 64, 125, 216, 343, 512, 729]
```

## Question 6:

Create a function that takes a string and counts the total number of capital letters in that string.



<!-- grade_id: comb_ans, points: 0, grade: false, solution: true -->
```python
def count_capital_letters(input_string):
    """Write a Python function to count the number of capital letters in a given string."""
    # YOUR CODE HERE
def count_capital_letters(input_string):
    """Count the number of capital letters in a given string."""
    return sum(1 for c in input_string if c.isupper())

# Test the function
print(count_capital_letters("Hello World!"))


```

Output:
```
2
```

```python
count_capital_letters('CE880!')
```

Output:
```
2
```

<!-- grade_id: comb_test, points: 1, grade: true, solution: false -->
```python
assert count_capital_letters('Essex University') == 2
assert count_capital_letters('Camulodunum the Ancient Roman') == 3
assert count_capital_letters('cHgYjHkYhUgYj12') == 6
```

```python

```